<a href="https://colab.research.google.com/github/stevendemer/machine-learning/blob/main/assignment2_24_25(1).ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Εργασία 2 ##

Καλωσήρθατε στην δεύτερη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα δενδρικά μοντέλα και την αξιολόγηση μοντέλων.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού **Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων**.

Συμπληρώστε το όνομα (NAME) και το AEM σας παρακάτω:

```python
NAME = "Stefanos Demertzis"
AEM = "3467"
```

**1 (0,5 μονάδες):** Διαβάστε το διαθέσιμο από την sklearn σύνολο δεδομένων diabetes που αφορά την πρόβλεψη της πορείας που θα έχει ένας ασθενής με διαβήτη ένα χρόνο μετά (εργασία παλινδρόμησης) και αποθηκεύστε το σε μια μεταβλητή με το όνομα diabetes.

<!-- grade_id: cell-945e6199220ebe51, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_diabetes
diabetes = load_diabetes()
```

<!-- grade_id: cell-5ad298e30b073224, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθής ανάγνωσης του συνόλου δεδομένων"""
assert diabetes.feature_names == ['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

**2 (0,5 μονάδες):** Χωρίστε το σύνολο που διαβάσατε σε υποσύνολο ελέγχου 20% και υποσύνολο εκπαίδευσης 80% με την μέθοδο train_test_split και σπόρο τυχαιότητας 42. Aποθηκεύστε το σύνολο εκπαίδευσης σε μεταβλητές X_train, y_train και το σύνολο ελέγχου σε μεταβλητές X_test, y_test. (0.5 μονάδες)

<!-- grade_id: cell-007e5ef64c2edaaf, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
X = diabetes.data
y = diabetes.target

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
```

<!-- grade_id: cell-33de7b4d5141b07e, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == 0.02736
assert round(X_test[0][8], 5) == 0.03243
```

**3 (1 μονάδα):** Χρησιμοποιώντας τα δεδομένα εκπαίδευσης, εκπαιδεύστε δύο δένδρα παλινδρόμησης με κριτήρια διαχωρισμού τα squared_eror και poisson ([δείτε ενότητα 1.10.7.2](https://scikit-learn.org/stable/modules/tree.html#tree)) αντίστοιχα, θέτοντας τον σπόρο τυχαιότητας στην τιμή 42 και το μέγιστο βάθος στην τιμή 3. Υπολογίστε τη μετρική r2 με βάση τις προβλέψεις τους στα δεδομένα ελέγχου, τις οποίες αποθηκεύστε στις μεταβλητές y_pred1 (squared_error) και y_pred2 (poisson) αντίστοιχα.
Αποθηκεύστε τις μετρικές στις μεταβλητές r2_squared_error και r2_poisson αντίστοιχα. (1 μονάδα)

<!-- grade_id: cell-b1390f68c6c3dc60, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import r2_score

tree_squared = DecisionTreeRegressor(criterion='squared_error', max_depth=3, random_state=42)
tree_squared.fit(X_train, y_train)

tree_poisson = DecisionTreeRegressor(criterion='poisson', max_depth=3, random_state=42)
tree_poisson.fit(X_train, y_train)

y_pred1 = tree_squared.predict(X_test)
y_pred2 = tree_poisson.predict(X_test)

r2_squared_error = r2_score(y_test, y_pred1)
r2_poisson = r2_score(y_test, y_pred2)
```

<!-- grade_id: cell-3f2d723addec3c0f, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού των μετρικών"""
assert round(r2_squared_error, 3) == 0.329
assert round(r2_poisson, 3) == 0.391
assert len(y_pred1) == len(y_pred2)
```

**4 (0,5 μονάδες):** Διαχωρίστε τα δεδομένα diabetes σε υποσύνολα εκπαίδευσης 60% (X_train, y_train), επικύρωσης 15% (X_val, y_val) και ελέγχου 25% (X_test, y_test) με σπόρο τυχαιότητας 42. Ξεκινήστε δημιουργώντας πρώτα το τελικό σύνολο ελέγχου, και έπειτα χωρίζοντας το υπόλοιπο σε σύνολο εκπαίδευσης και επικύρωσης.

<!-- grade_id: cell-db6d2d4c83a94e4b, points: 0, grade: false, solution: true -->
```python
X_temp, X_test, y_temp, y_test = train_test_split(
    X, y,
    test_size=0.25,
    random_state=42
)
X_train, X_val, y_train, y_val = train_test_split(
    X_temp, y_temp,
    test_size=0.2,
    random_state=42
)

```

<!-- grade_id: cell-262a292a0b3f72bb, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == -0.08024
assert round(X_val[0][8], 5) == 0.09924
assert round(X_test[0][8], 5) == 0.03243
```

**5 (2 μονάδες):** Υλοποιήστε συνάρτηση, η οποία δοθέντος ενός δενδρικού μοντέλου παλινδρόμησης, διερευνά όλες τις ακέραιες τιμές της μεταβλητής min_samples_leaf από ένα κάτω όριο μέχρι ένα άνω όριο και επιστρέφει την βέλτιστη με βάση το mean squared error στο παραπάνω σύνολο επικύρωσης έπειτα από εκπαίδευση στο παραπάνω σύνολο εκπαίδευσης. Επιπλέον επιστρέφει και το mean squared error (στρογγυλοποιημένο σε ακέραιο) που πετυχαίνει στο σύνολο ελέγχου το μοντέλο με την βέλτιστη επίδοση στο σύνολο επικύρωσης.

<!-- grade_id: cell-db6235d877cf3abe, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error
import numpy as np

def best_min_samples_leaf(model, min, max):
   best_mse = float('inf')
   best_leaf = None

   for min_leaf in range(min, max + 1):
        current_model = DecisionTreeRegressor(
            criterion=model.criterion,
            max_depth=model.max_depth,
            min_samples_leaf=min_leaf,
            random_state=model.random_state
        )
        current_model.fit(X_train, y_train)


        val_pred = current_model.predict(X_val)
        current_mse = mean_squared_error(y_val, val_pred)


        if current_mse < best_mse:
            best_mse = current_mse
            best_leaf = min_leaf
            best_model = current_model


   test_pred = best_model.predict(X_test)
   test_mse = int(round(mean_squared_error(y_test, test_pred)))

   return best_leaf, test_mse



```

<!-- grade_id: cell-4c4b39c4be673d88, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
assert best_min_samples_leaf(model, 1, 8) == (7, 3358)
assert best_min_samples_leaf(model, 5, 15) == (11, 3271)
```

**6 (1,5 μονάδα):** Υλοποιήστε συνάρτηση η οποία δοθέντος ενός μοντέλου, ενός συνόλου δεδομένων και μίας μετρικής, αξιολογεί το μοντέλο με τη μέθοδο της σταυρωτής επικύρωσης "άφησε ένα εκτός".

<!-- grade_id: cell-83aa7655f3f1ecdd, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import LeaveOneOut
import numpy as np
from sklearn.metrics import get_scorer

def leave_one_out(model, X, y, metric):
    loo = LeaveOneOut()
    scores = []

    for train_index, test_index in loo.split(X):
        X_train, X_test = X[train_index], X[test_index]
        y_train, y_test = y[train_index], y[test_index]

        model.fit(X_train, y_train)
        y_pred = model.predict(X_test)

        if metric == mean_squared_error:
          score = metric(y_test, y_pred)
        else:
          score = get_scorer(metric)(model, X_test, y_test)


        scores.append(score)

    return np.mean(scores)




```

<!-- grade_id: cell-06bac2fcee9a42de, points: 1.5, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_leaf_nodes=6, min_samples_leaf=2, random_state=42)
assert round(leave_one_out(model, diabetes.data, diabetes.target, mean_squared_error), 0) == 3354
```

**7 (2 μονάδες):** Χρησιμοποιώντας τις κλάσεις GridSearchCV, KFold και τη συνάρτηση cross_val_score, υλοποιήστε συνάρτηση που δέχεται δενδρικό μοντέλο, σύνολο δεδομένων, μετρική αξιολόγησης και ελάχιστη (min) και μέγιστη (max) τιμή βάθους για το δένδρο και επιστρέφει εκτίμηση της απόδοσης του μοντέλου αυτού, έπειτα από επιλογή της καλύτερης δυνατής τιμής για την παράμετρο του βάθους από min έως και max, με βάση τη δοθείσα μετρική και τη μέθοδο της φωλιασμένης k-πλης σταυρωτής επικύρωσης. Χρησιμοποιήστε 10 fold εξωτερικά, 9 fold εσωτερικά, τυχαίο ανακάτεμα δεδομένων, και σπόρο τυχαιότητας 25.

<!-- grade_id: cell-0859aa898e82771f, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import GridSearchCV, KFold, cross_val_score
import numpy as np

def estimation(model, X, y, metric, min, max):
  # Define the parameter grid for max_depth
    param_grid = {'max_depth': list(range(min, max + 1))}

    # Set up cross-validation strategies
    inner_cv = KFold(n_splits=9, shuffle=True, random_state=25)
    outer_cv = KFold(n_splits=10, shuffle=True, random_state=25)

    # Set up grid search for hyperparameter tuning
    grid_search = GridSearchCV(
        estimator=model,
        param_grid=param_grid,
        cv=inner_cv,
        scoring=metric,
        n_jobs=-1  # Use all available cores for faster computation
    )

    # Perform nested cross-validation
    nested_scores = cross_val_score(
        grid_search,
        X=X,
        y=y,
        cv=outer_cv,
        scoring=metric,
        n_jobs=-1
    )

    # Return the mean score rounded to 2 decimal places
    return np.mean(nested_scores)


```

<!-- grade_id: cell-29fa8f9d02b1331e, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(random_state=42)
assert round(estimation(model, diabetes.data, diabetes.target, 'neg_mean_squared_error', 1, 10), 2) == 3993.82
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-38-a157e56a3aa6> in <cell line: 0>()
      1 """Τεστ ορθής υλοποίησης συνάρτησης"""
      2 model = DecisionTreeRegressor(random_state=42)
----> 3 assert round(estimation(model, diabetes.data, diabetes.target, 'neg_mean_squared_error', 1, 10), 2) == 3993.82

AssertionError:
```

**8 (1 μονάδα):** Υλοποιήστε συνάρτηση που δέχεται εκπαιδευμένο δενδρικό μοντέλο και πίνακα συμβολοσειρών με τα ονόματα των μεταβλητών εισόδου και επιστρέφει τη σημαντικότερη μεταβλητή εισόδου.

<!-- grade_id: cell-df6dc2f47f9d098b, points: 0, grade: false, solution: true -->
```python
def best_feature(model, feature_names):
    if not hasattr(model, 'feature_importances_'):
        raise ValueError("Το μοντέλο δεν έχει υποστεί εκπαίδευση ή δεν υποστηρίζει feature importances")


    if len(feature_names) != len(model.feature_importances_):
        raise ValueError("Το μήκος των feature names δεν ταιριάζει με το πλήθος των features του μοντέλου")


    most_important_idx = model.feature_importances_.argmax()


    return feature_names[most_important_idx]



```

<!-- grade_id: cell-50285e47d84e6828, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
model.fit(diabetes.data, diabetes.target)
assert best_feature(model, diabetes.feature_names) == 's5'
```

**9 (1 μονάδα):** Υλοποιήστε συνάρτηση που δέχεται εκπαιδευμένο δενδρικό μοντέλο και πίνακα συμβολοσειρών με τα ονόματα των μεταβλητών και επιστρέφει την μεταβλητή η οποία ελέγχεται στο δένδρο στο πρώτο παιδί δεξιά. Χρησιμοποιήστε την ιδιότητα [tree_](https://scikit-learn.org/stable/auto_examples/tree/plot_unveil_tree_structure.html#sphx-glr-auto-examples-tree-plot-unveil-tree-structure-py).

<!-- grade_id: cell-ead77d5ea43abcf1, points: 0, grade: false, solution: true -->
```python
def right_child(model, feature_names):
    if not hasattr(model, 'tree_'):
        raise ValueError("Το μοντέλο δεν έχει εκπαιδευτεί")

    tree = model.tree_


    if tree.children_left[0] == -1 and tree.children_right[0] == -1:
        return None


    right_child = tree.children_right[0]


    feature_idx = tree.feature[right_child]


    if feature_idx == -2:
        return None

    return feature_names[feature_idx]

```

<!-- grade_id: cell-c59ba7d856846a2c, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=2, random_state=42)
model.fit(diabetes.data, diabetes.target)
assert right_child(model, diabetes.feature_names) == 'bmi'
```